# HW1: Các câu lệnh Linux, Python

(Cập nhật lần cuối: 30/10/2021)

Họ tên: Hồ Hoàng Duy

MSSV: 19120207

---

## Cách làm bài và nộp bài

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Bạn cố gắng giữ cho code rõ ràng, ngăn nắp: đặt tên biến gợi ý, comment ở những chỗ cần comment, giữ cho một dòng code không quá dài (vì nếu không thì người đọc sẽ phải mất công scroll, bạn có thể dùng qui tắc: một dòng code chỉ nên có tối đa là 79 ký tự). 

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart & Run All` để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart & Run All` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW1.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài.</font>

---

## Môi trường code

Ta thống nhất trong môn này:
- Code trên môi trường Linux (bạn đã setup ở HW0)
- Dùng Jupyter Notebook hoặc Jupyter Lab để code file .ipynb. Lưu ý: bạn <font color=red>không dùng Google Colab</font>, vì Google Colab có thể xóa một số thông tin meta mà chương trình hỗ trợ chấm bài của mình đã tạo trong file bài tập
- Dùng phiên bản các package như trong file "min_ds-env.yml" (bạn xem trong HW0 cách dùng Conda để tạo môi trường code từ file .yml và cách kích hoạt môi trường này sau khi đã tạo)

```python
import sys
sys.executable
```

Output:
```
'/home/hhduy/miniconda3/envs/min_ds-env/bin/python'
```

Nếu không có vấn đề gì thì file chạy python sẽ là file của môi trường code "min_ds-env".

Trong bài này, ngoài Git & Jupyter Notebook/Lab thì bạn chỉ dùng Python, không dùng các thư viện khác như Numpy, Pandas, ... Trong Python, khi đọc/ghi file csv thì bạn cũng không được dùng các thư viện hỗ trợ như thư viện `csv`.

---

## Import các thư viện cần thiết

```python
import datetime
# YOUR CODE HERE (OPTION)
import os.path
import subprocess
```

---

## Thu thập dữ liệu (3.5đ)

Trong bài tập này, bạn sẽ lấy dữ liệu log của [thùng chứa nbdime](https://github.com/KienTrann/nbdime). Nbdime là công cụ giúp Git có thể diff và merge file Notebook một cách phù hợp (thay vì chỉ coi như là một file text thông thường). Link mình cung cấp không phải là link đến thùng chứa gốc của nbdime mà là thùng chứa mình "fork" (đại khái là tạo ra một bản copy) từ thùng chứa gốc. Lý do: mình muốn cố định lại dữ liệu log để có thể cung cấp các bộ test để bạn phần nào biết là đã code đúng hay chưa (còn nếu dùng thùng chứa gốc thì dữ liệu log có thể sẽ thay đổi trong vài ngày tới, các bộ test lúc đó có thể sẽ không còn đúng nữa).

Để có thể lấy dữ liệu log thì việc đầu tiên bạn cần làm ở cell bên dưới là: 

- Dùng câu lệnh Git để clone thùng chứa về; bạn clone vào một thư mục cùng cấp với file notebook này. <font color=gray>Gợi ý: trong Jupyter Notebook, ở code cell, bạn có thể thực hiện một câu lệnh Linux (bao gồm câu lệnh Git) bằng cách thêm dấu ! ở đầu.</font>
- Sau đó, để xem thử thư mục vừa clone có nặng không, bạn tính dung lượng tổng cộng của thư mục này và lưu vào biến `repo_size` (`repo_size` là một con số, đơn vị là MB); ta thống nhất: (i) dùng cờ `apparent-size` khi dùng câu lệnh `du`, (ii) code chỉ cần chạy đúng với thùng chứa cụ thể này. <font color=gray>Gợi ý: trong Jupyter Notebook, bạn có thể gán kết quả của câu lệnh Linux vào biến Python.</font>

<!-- grade_id: cell-5984c6f1d3bb0311, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
if (os.path.isdir('/nbdime') == False):
    !git clone https://github.com/KienTrann/nbdime.git
# Thầy HDTH với mình có test thử thì có máy ra 9.7, có máy ra 9.8
# (có thể là do phiên bản của câu lệnh du?)
repo_size = subprocess.run(['du','-sh', '--apparent-size', 'nbdime']).stdout
```

Output:
```
fatal: destination path 'nbdime' already exists and is not an empty directory.
9,8M	nbdime
```

Sau khi đã clone thùng chứa về, việc kế tiếp của bạn là dùng các câu lệnh Linux (trong đó có các câu lệnh Git) và Python để lấy dữ liệu log rồi lưu xuống file "log.tsv" (định dạng tsv tương tự như csv nhưng các giá trị được cách nhau bởi tab). Để dễ hình dung, bạn có thể dùng Jupyter Lab để mở xem file kết quả đúng "correct_log.tsv" mà mình cung cấp. Như bạn có thể thấy, mỗi dòng ứng với một commit (không tính dòng đầu - dòng chứa các tên cột) và gồm các thông tin:
- "Id" là id (hash) dạng viết gọn của commit
- "ParentIds" là id dạng viết gọn của (các) commit cha
- "AuthorName" là tên của tác giả
- "AuthorEmail" là email của tác giả
- "AuthorDate" là ngày tác giả thực hiện commit, ngày được format theo chuẩn RFC2822 (xem phần gợi ý ở bên dưới để biết cách lấy ngày theo format này) 
- "Subject" là tựa đề của message (message tương tự như email, ngoài phần "subject" thì có thể có thêm phần "body"; ở đây ta chỉ lấy phần "subject")
- "ChangedFiles": các file bị thay đổi, nếu có nhiều file thì các file được phân tách bởi dấu phẩy

Tại sao mình dùng định dạng tsv thay vì csv? Vì dấu phẩy nhiều khả năng là sẽ có trong phần "subject" của message. Tab ít khả năng sẽ có hơn, nhưng nếu có thì bạn hãy thay bằng bốn khoảng trắng.

<font color=gray>Gợi ý cho phần này:</font> 
- <font color=gray>Bạn có thể xem thêm về câu lệnh `git log` [ở đây](https://git-scm.com/book/en/v2/Git-Basics-Viewing-the-Commit-History)</font>
- <font color=gray>Khi ở Notebook, bạn thử `cd` đến thư mục clone rồi `pwd` xem? Kết quả như vậy là vì khi thực hiện một dòng lệnh Linux ở Notebook thì Notebook sẽ tạo ra một terminal tạm thời để thực thi dòng lệnh đó trong phút chốc. Như vậy, câu lệnh `cd` vẫn là có `cd` đến thư mục clone nhưng là ở terminal tạm thời; khi câu lệnh `cd` chạy xong thì terminal tạm thời này sẽ được đóng và ta vẫn ở thư mục mà hiện chứa file Notebook. Để có thể `cd` rồi `pwd` thì một cách là ghi 2 câu lệnh này trên một dòng và cách nhau bởi `&&`; lúc này thì cả 2 câu lệnh sẽ được thực thi trên cùng một terminal</font>

<!-- grade_id: cell-7a31931ee1dd0b8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
!cd nbdime && pwd && git log --oneline --name-only --pretty=format:%h,%p,%an,%ae,%aD,%s > log.tsv && mv -f log.tsv ../
#--pretty=format:"%h%x09%p%x09%an%x09%ae%x09%aD%x09%s%x09" 
f = open('log.tsv', 'r')
line_data = f.readlines()
file_changed = []
infor = []
#nếu dòng dữ liệu có nhiều thông tin phân cách nhau bởi dấu "," thì thêm vào list infor, rồi sau đó gán lại = 0 
for i in range(len(line_data)):
    if (',' in line_data[i]):
        infor.append(line_data[i])
        line_data[i] = 0
#nếu dòng thứ i = 0, ta xét dòng j = i + 1 trở về sau
for i in range(len(line_data)):
    if (line_data[i] == 0):
        file_modified = [] #list to store modified files
        count_file_modified = 0 #to store number of files is modified in one commit
        for j in range(i + 1, len(line_data)):
            if (line_data[j] == 0):
                break
            else:
                file_modified.append(line_data[j].replace('\n', '')) #append thông tin dòng (chính là file đã modified)
                count_file_modified += 1
        if (count_file_modified == 0):
            file_changed.append(0)
        else:
            file_changed.append(file_modified)
Id = []
ParentIds = []
AuthorName = []
AuthorEmail = []
AuthorDate = []
Subject = []
ChangedFiles = []
for info in infor:
    inf = info.split(',')
    Id.append(inf[0])
    ParentIds.append(inf[1])
    AuthorName.append(inf[2])
    AuthorEmail.append(inf[3])
    #khi ta tách dòng thông tin thì định dạng ngày theo RFC2822 cũng bị tách ra, để hiển thị lại thì ta phải 
    #ghép chuỗi inf4 với inf5 lại
    date = inf[4] + ',' + inf[5]
    AuthorDate.append(date)
    temp = inf[6]
    if (len(inf) == 7): #nếu dòng có 7 thông tin
        temp = temp.replace('\n', '')
        Subject.append(temp)
    else: #nếu dòng có nhiều hơn 7 thông tin => ghép chuỗi subject lại (vì message có dấu phẩy)
        for j in range(7, len(inf)):
            temp = temp + ',' + inf[j]
        temp = temp.replace('\n', '')
        Subject.append(temp)
for i in file_changed:
    if(i == 0):
        ChangedFiles.append('\t')
    else:
        fileName = ','.join(i)
        fileName.replace('\n', '')
        fileName.replace('\t', '')
        if(fileName[-1] == ',' or fileName[-1] == ' '):
            ChangedFiles.append(fileName[:-1])
        else:
            ChangedFiles.append(fileName + ' ')
with open('log.tsv', 'w+') as f:
    f.write("Id\tParentIds\tAuthorName\tAuthorEmail\tAuthorDate\tSubject\tChangedFiles\n")
    for i in range(len(Id)):
        f.write(Id[i] + '\t' + ParentIds[i] + '\t' + AuthorName[i] + '\t' + AuthorEmail[i] + '\t' + AuthorDate[i] + '\t' + Subject[i] + '\t')
        #nếu có file được thay đổi => viết thêm vào sau \tab
        if(ChangedFiles[i] != '\t'):
            if(i == len(Id) - 1):
                f.write(ChangedFiles[i][:-1])
            else:
                f.write(ChangedFiles[i])
        f.write('\n')
    f.close()
```

Output:
```
/home/hhduy/Documents/P4DS/HW1/nbdime
```

<!-- grade_id: cell-1671c3f106d5da7d, points: 3, grade: true, solution: false -->
```python
# TEST
file = open('log.tsv', 'r')
log = file.read()
file.close()
file = open('correct_log.tsv', 'r')
correct_log = file.read()
file.close()
assert log == correct_log
```

---

## Khám phá dữ liệu

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.5đ)

Đầu tiên, bạn sẽ đọc dữ liệu từ file "log.tsv" (nếu phần thu thập dữ liệu bạn làm chưa đúng thì bạn có thể đọc từ file "correct_log.tsv" mà mình đã cung cấp). Bạn lưu kết quả đọc được vào biến `cols`. Biến `cols` là một dictionary, trong đó mỗi phần tử ứng với một cột, key là chuỗi ứng với tên cột, value là list chứa các giá trị của cột (mỗi giá trị bạn cứ tạm để kiểu dữ liệu là chuỗi).

Kế tiếp, từ biến `cols`, bạn sẽ tính số dòng và số cột, rồi lần lượt lưu vào hai biến là `num_rows` và `num_cols`.

<!-- grade_id: cell-f2f6dbf283e0ccee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
read_data = []
with open('log.tsv', encoding='utf-8') as file:
    for line in file:
        li = line.split('\t')
        read_data.append(li)
read_data.pop(0) #bỏ phần tử đầu tiên (chứa list tên thuộc tính)
Id = []
ParentIds = []
AuthorName = []
AuthorEmail = []
AuthorDate = []
Subject = []
ChangedFiles = []
for inf in read_data:
    Id.append(inf[0])
    ParentIds.append(inf[1])
    AuthorName.append(inf[2])
    AuthorEmail.append(inf[3])
    AuthorDate.append(inf[4])
    Subject.append(inf[5])
    ChangedFiles.append(inf[6].replace('\n', ''))
cols = {'Id': Id, 'ParentIds': ParentIds, 'AuthorName': AuthorName, 'AuthorEmail': AuthorEmail, 'AuthorDate': AuthorDate, 'Subject': Subject, 'ChangedFiles': ChangedFiles}
num_rows = len(Id)
num_cols = len(list(cols.keys()))
```

<!-- grade_id: cell-cae77558e0985b8d, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert num_rows == 1928
assert num_cols == 7
assert set(cols.keys()) == set(['Id', 'ParentIds', 'AuthorName', 
                                'AuthorEmail', 'AuthorDate', 'Subject', 
                                'ChangedFiles'])
id_col_len = len(cols['Id'])
for col_name in cols:
    assert len(cols[col_name]) == id_col_len
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Mỗi dòng ứng với một commit. Có vẻ không có dòng nào "lạc loài" (nghĩa là không có vấn đề các dòng có ý nghĩa khác nhau).

### Dữ liệu có các dòng bị lặp không? (0.25đ)

Có lẽ là không có. Để chắc ăn thì một cách là ta có thể check cột "Id"; nếu cột "Id" không có giá trị bị lặp thì có thể kết luận là không có dòng bị lặp. Bạn sẽ làm vụ này và lưu kết quả check vào biến `duplicated_id`; biến này bằng True nếu cột "Id" có giá trị lặp và bằng False nếu ngược lại.

<!-- grade_id: cell-31e053e964e6c9eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
check_Id = cols['Id']
duplicated_id = False
for i in range(1, len(check_Id)):
    for j in range(i):
        if check_Id[i] == check_Id[j]:
            duplicated_id = True
```

<!-- grade_id: cell-ab5957518b489fa7, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert duplicated_id == False
```

### Mỗi cột có ý nghĩa gì?

Vụ này mình đã nói lúc thu thập dữ liệu. Ở đây, mình ghi lại một lần nữa:

- "Id" là id (hash) dạng viết gọn của commit
- "ParentIds" là id dạng viết gọn của (các) commit cha
- "AuthorName" là tên của tác giả
- "AuthorEmail" là email của tác giả
- "AuthorDate" là ngày tác giả thực hiện commit
- "Subject" là tựa đề của message (message tương tự như email, ngoài phần "subject" thì có thể có thêm phần "body"; ở đây ta chỉ lấy phần "subject")
- "ChangedFiles": các file bị thay đổi, nếu có nhiều file thì các file được phân tách bởi dấu phẩy

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không?

- Các cột "Id", "ParentIds", "AuthorName", "AuthorEmail", "Subject": kiểu dữ liệu dạng chuỗi là đúng rồi
- Cột "ChangedFiles": hmm ... có lẽ tạm thời ta vẫn để dạng chuỗi, có gì cần thì ta sẽ xử lý sau
- Cột "AuthorDate": cần chuyển từ kiểu dữ liệu dạng chuỗi sang datetime; để có thể tiếp tục khám phá tiếp thì ngay sau đây bạn sẽ tiền xử lý vụ này luôn

---

## Tiền xử lý (0.5đ)

Ở đây, bạn sẽ chuyển cột "AuthorDate" từ dạng chuỗi sang datetime. <font color=gray>Gợi ý: bạn xem [document](https://docs.python.org/3/library/datetime.html#strftime-and-strptime-format-codes) để biết cách truyền format phù hợp cho hàm `strptime`.</font>

<!-- grade_id: cell-742e74182a9d1936, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
for i in range(num_rows):
    fmt_date = cols['AuthorDate'][i][:-6]
    fmt_date = datetime.datetime.strptime(fmt_date, "%a, %d %b %Y %X")
    fmt_date = fmt_date.replace(tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
    cols['AuthorDate'][i] = fmt_date
```

<!-- grade_id: cell-c551b6a15162bb10, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert type(cols['AuthorDate'][0]) == datetime.datetime
assert cols['AuthorDate'][0] == datetime.datetime(2021, 4, 20, 12, 27, 9, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
```

---

## Khám phá dữ liệu (tiếp tục)

### Với mỗi cột có kiểu dữ liệu dạng số (numeric) và dạng datetime, các giá trị được phân bố như thế nào? (0.5đ)

Ở đây thì không có cột dạng số, chỉ có một cột dạng datetime là cột "AuthorDate".

Đầu tiên, bạn tính tỉ lệ % các giá trị thiếu của cột "AuthorDate". <font color=gray>Gợi ý: không cần code, có thể lập luận dựa vào kết quả khi chuyển cột "AuthorDate" từ dạng chuỗi sang datetime ở bên trên.</font>

<!-- grade_id: cell-47e2584daf9dda74, points: 0.25, grade: true, solution: true -->
### YOUR ANSWER HERE
Dựa vào kết quả phần TEST khi tiền xử lý cột 'AuthorDate' thì ta có thể biết được tỉ lệ % các giá trị thiếu của cột này là 0% (có nghĩa là không có giá trị nào bị thiếu trong cột cả).

Tiếp theo, bạn tính min và max của cột "AuthorDate", rồi lần lượt lưu vào hai biến `author_date_min` và `author_date_max`.

<!-- grade_id: cell-ed5647763d70e921, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
author_date = cols['AuthorDate']
author_date_min = min(author_date)
author_date_max = max(author_date)
```

<!-- grade_id: cell-8bb55737a75c0a6e, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert author_date_min == datetime.datetime(2015, 11, 9, 13, 13, 28, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
assert author_date_max == datetime.datetime(2021, 4, 20, 12, 27, 9, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
```

Hmm ... Min và max của cột "AuthorDate" không có gì bất thường.

### Với mỗi cột có kiểu dữ liệu dạng phân loại (categorical), các giá trị được phân bố như thế nào? (1.75đ)

Các cột dạng phân loại là các cột mà không phải là cột "AuthorDate". Bạn sẽ lưu thông tin của các cột này vào biến `cate_col_profiles`. Biến này là một dictionary, trong đó mỗi phần tử ứng với một cột dạng phân loại, key là tên cột, value là list gồm: tỉ lệ % giá trị thiếu, số lượng các giá trị khác nhau (không tính giá trị thiếu), set chứa các giá trị khác nhau (không bao gồm giá trị thiếu).

<!-- grade_id: cell-d9eb895996d7b157, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
cate_col_profiles = {}
author_email = []
author_name = []
changed_files = []
Id = []
parent_ids = []
subject = []
key = ['AuthorEmail', 'AuthorName', 'ChangedFiles', 'Id', 'ParentIds', 'Subject']
for k in key:
    # set của mỗi khóa k là set của cột cols (cols[k])
    k_val = set(cols[k])
    # nếu chuỗi rỗng hoặc \n thì xóa khỏi tập k_val
    if ('' in k_val):
        k_val.remove('')
    elif ('\n' in k_val):
        k_val.remove('\n')
    cnt = 0 # biến đếm sốgiá trị thiếu
    for val in cols[k]:
        if(val not in k_val):
            cnt += 1
    cate_col_profiles[k] = [cnt * 100 / len(cols[k]), len(k_val), k_val]
#cate_col_profiles
```

<!-- grade_id: cell-50123f76f79e1f27, points: 1, grade: true, solution: false -->
```python
# TEST
assert set(cate_col_profiles.keys()) == {'AuthorEmail', 'AuthorName', 
                                         'ChangedFiles', 'Id', 
                                         'ParentIds', 'Subject'}
for col_name, col_profile in cate_col_profiles.items():
    assert len(col_profile) == 3
    assert len(col_profile[2]) == col_profile[1]
assert round(cate_col_profiles['ParentIds'][0], 3) == 0.052
assert round(cate_col_profiles['ChangedFiles'][0], 3) == 16.079
assert cate_col_profiles['ParentIds'][1] == 1820
assert cate_col_profiles['ChangedFiles'][1] == 842
```

```python
# TEST (BY EYES)
print(f"{'ColName':12} {'Miss(%)':7} {'NumDifVals':10} {'SomeVals'}")
for col_name, col_profile in cate_col_profiles.items():
    print(f'{col_name:12} {col_profile[0]:<7.3f} {col_profile[1]:<10} {col_profile[2].__repr__()[:34]+"...":<}')
```

Output:
```
ColName      Miss(%) NumDifVals SomeVals
AuthorEmail  0.000   34         {'tony.hirst@gmail.com', 'piotr.ch...
AuthorName   0.000   34         {'Will Engler', 'Alex Bozarth', 'T...
ChangedFiles 16.079  842        {'nbdime/webapp/src/app/cm-merge.c...
Id           0.000   1928       {'bcc1734', '5d2a9f4', '3fefdc8', ...
ParentIds    0.052   1820       {'bcc1734', '5d2a9f4', 'd15e6fd 8e...
Subject      0.000   1860       {'Show README content as package d...
```

Bạn quan sát kết quả in ra ở trên (tạm thời chưa để ý đến cột "Miss(%)") và cho biết có gì bất thường không? Hay mọi thứ đều bình thường? Cho mình biết tại sao bạn lại kết luận như vậy.

<!-- grade_id: cell-88055a6ed4991805, points: 0.25, grade: true, solution: true -->
### YOUR ANSWER HERE


Kế tiếp, bạn hãy cho in ra các dòng mà có giá trị thiếu (in sao là tùy bạn, miễn là dễ nhìn). Sau đó, bạn hãy quan sát các dòng này để xem tại sao lại bị thiếu (trong trường hợp, bạn không biết tại sao thì cứ ghi nhận là không biết tại sao).

<!-- grade_id: cell-fda56cd71a2bb407, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
/tmp/ipykernel_37574/1067236896.py in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-0ce04cd6c1dd7a31, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Đưa ra câu hỏi có ý nghĩa cần trả lời

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có ý nghĩa mà có thể được trả lời bằng dữ liệu này không.

Các câu hỏi có thể có là: 
- *Câu hỏi 1: trong năm 2021, ai là người có nhiều commit nhất, ai là người có nhiều commit nhì, ...?* \
Tìm câu trả lời cho câu hỏi này để làm gì? \
Hmm ... Chẳng hạn để khen thưởng
- *Câu hỏi 2: commit nào là commit cuối cùng đã thay đổi file "nbdime/webapp/templates/difftool.html"?* \
Tìm câu trả lời cho câu hỏi này để làm gì? \
Hmm ... Chẳng hạn để đổ tội cho tác giả của commit đó ;-)

---

## Tiền xử lý (nếu cần) + phân tích dữ liệu để trả lời cho câu hỏi 1 (1đ)

Nhắc lại câu hỏi 1: *trong năm 2021, ai là người có nhiều commit nhất, ai là người có nhiều commit nhì, ...?*

Bạn sẽ lưu kết quả vào biến `wanted_people`. Biến này là một dictionary, mỗi phần tử ứng với một người mà có commit trong năm 2021, key là tên của người đó, value là số lượng commit của người đó. Dictionary này được sort theo value giảm dần.

<!-- grade_id: cell-a40eed2ac02f6c2b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
wanted_people = {}
start_date = datetime.datetime(2021, 1, 1 ,tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
end_date = datetime.datetime(2021, 12, 31 ,tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
for i in range(num_rows):
    cnt = 0
    # nếu tác giả đã có trong dict wanted_people thì tiếp tục đếm
    if (cols['AuthorName'][i] in wanted_people.keys()):
        continue
    for j in range(num_rows):
        if (cols['AuthorName'][j] == cols['AuthorName'][i]) and (cols["AuthorDate"][j].year == 2021):
            cnt += 1
    wanted_people[cols["AuthorName"][i]] = cnt
for i in list(wanted_people.keys()):
    if(wanted_people[i] == 0):
        wanted_people.pop(i) #bỏ đi tên những tác giả có số commit = 0
# {k: v for k, v in sorted(x.items(), key=lambda item: item[1])}
# dict(sorted(x.items(), key=lambda item: item[1]))
wanted_people = {k: v for k, v in sorted(wanted_people.items(), key=lambda item: item[1], reverse=True)}
wanted_people
```

Output:
```
{'Vidar Tonaas Fauske': 26,
 'Alex Bozarth': 9,
 'Frederic COLLONVAL': 7,
 'Frédéric Collonval': 6,
 'krassowski': 2,
 'Michał Krassowski': 2}
```

<!-- grade_id: cell-9e312d16b92a0e40, points: 1, grade: true, solution: false -->
```python
# TEST
correct_wanted_people = {'Vidar Tonaas Fauske': 26, 
                         'Alex Bozarth': 9, 
                         'Frederic COLLONVAL': 7, 
                         'Frédéric Collonval': 6, 
                         'krassowski': 2, 
                         'Michał Krassowski': 2}
assert wanted_people == correct_wanted_people
```

---

## Tiền xử lý (nếu cần) + phân tích dữ liệu để trả lời cho câu hỏi 2 (2đ)

Nhắc lại câu hỏi 2: *commit nào là commit cuối cùng đã thay đổi file "nbdime/webapp/templates/difftool.html"?*

Ta thống nhất: không xét các commit "merge pull request" vì các commit này không có thông tin về "ChangedFiles".

Bạn lưu kết quả vào biến `wanted_commit`. Biến này là một dictionary, trong đó mỗi phần tử ứng với một thông tin của commit, key là tên cột, value là giá trị tương ứng.

<!-- grade_id: cell-37a6a90f597ad4a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41c49732faf571f7, points: 2, grade: true, solution: false -->
```python
# TEST
correct_wanted_commit = {\
    'Id': 'be0b5ee',
    'ParentIds': '4226179',
    'AuthorName': 'Vidar Tonaas Fauske',
    'AuthorEmail': 'vidartf@gmail.com',
    'AuthorDate': datetime.datetime(2017, 7, 4, 22, 53, 9, tzinfo=datetime.timezone(datetime.timedelta(seconds=7200))),
    'Subject': 'jinja2 templates + avoid name collisions with nb',
    'ChangedFiles': 'nbdime/webapp/nbdimeserver.py,nbdime/webapp/templates/compare.html,nbdime/webapp/templates/diff.html,nbdime/webapp/templates/difftool.html,nbdime/webapp/templates/merge.html,nbdime/webapp/templates/mergetool.html,nbdime/webapp/templates/nbdimepage.html,setup.py'\
}
assert wanted_commit == correct_wanted_commit
```

```python

```